```python
%%capture
# Install the autograder to check your answers as you go along
import sys
!{sys.executable} -m pip install otter-grader
# Initialize Otter
import otter
grader = otter.Notebook("HW2.ipynb")
```

# Homework #2
**Due Friday Oct. 27th @ 11:59 pm**

**Objective:** This assignment will give you experience with control flow, Numpy arrays and functions, datetimes, and basic plotting.


## **Instructions:**
### Accessing Class Code

1. Clone this repository into your own JupyterHub by running this command in your home directory:
```git clone your_SSH_URL```

2. Once you have cloned the repository, go into your directory and set your branch to "main" (See the GitHub cheatsheet for help on this). You will not need to reset your remote origin, as you have already directly cloned from your own version of the repository.

3. There should now be a "homework_2" directory in the home directory of your JupyterHub. In terminal, change directories into "homework_2". Next, click on the the "homework_1" icon on the filepath hierarchy in the left panel of JupyterHub. If you don't see it, make sure you're in the home folder by clicking the folder icon under the search bar. 

4. Double click the "HW2.ipynb" to open it in a new tab and begin working on the assignment. Read the instructions carefully, and make sure to write your answers in the specified cells. Typically, you will see a "..." in the places you need to fill in. Make sure to use the variable names provided in the starter code. See the "Working in your Notebook" section below for an example. There are some autograder tests embedded in the notebook, but there are also some hidden tests that will be graded after submission.

5. Edit the README file and write your name and UW NetID. Include a paragraph of instructions on how you used git commands to add, commit, and push to your repository (3-5 sentences). Now that y'all are GitHub connoisseurs, imagine that you are explaining how to use GitHub to someone who has never used GitHub before! :) 

6. As you continue to answer the homework questions and make edits to your code, make sure to regularly update your GitHub repository as well via git add, commit, and push (steps 15, 16, 19 in 0b). A good rule of thumb would be to run these git steps anytime you make an addition or change that you don't want to accidentally lose. Generally, you can push once a day to maintain good version control practices. <br>
As a note, make sure that your git commands are running without errors before you refresh your GitHub and check your changes. If you are not seeing the updated changes you created in your local JupyterHub directory, check where your status is by this command: <br>
``` git status```

Then, you can see if you made an error with your git add, commit, or push commands. 

**Sometimes, our JupyterHub server has trouble remembering the file permissions for our SSH keys. If you get a file permission error with your private ssh key, run this line of code:**
```chmod 400 ~/.ssh/id_ed25519``` 
<br>

This will change your file permission to the proper permissions that SSH requires.
    
### Submitting to GradeScope

7. Go to the class Gradescope dashboard and submit your personal GitHub repository link to the Homework 2 assignment. Make sure your GitHub is synchronized with Gradescope to access both your public _and_ private repositories. If prompted, log in to GitHub.

8. Run the autograder to check if your code runs and if you passed the initial unit tests. You should be able to run the autograder as many times as you want before submitting. Again, double check that your final answers are stored in the provided variable names given in the starter code!

13. Once the autograder has finished running, check that you have submitted the assignment. If you make any more changes to your code after submitting to Gradescope, make sure to push your changes to GitHub and resubmit the assignment on Gradescope. You can submit as many times as you want as there is no maximum submission attempts, but be sure to have your final submission in before the deadline.

### Working in your Notebook
   To help you start thinking about how to write meaningful and concise variable names, we have provided variable names in most of your questions. 
    Note that there is an ellipsis (the "...") after each of the variable names. These are the sections you are expected to fill in. Please use the provided variable names (ie "pelagic" and "coastal" in the above example) to report your final answer back in. This will ensure that the autograder on Gradescope runs properly. <br>
    <br>Make sure that you are adding comments and your outside references as you go along! Part of your grade will include using best coding practices in your homework assignments. 

### Honor Code

- Complete the assignment by writing and executing text and code cells as specified. For this assignment, do not use any features of Python that have not yet been discussed in the lessons or class sessions.

- Please keep in mind our late work and dropped homework grading policy. Review the syllabus for details.

- You can acknowledge and describe any assistance you've received on this assignment in the specified cell of the HW1 notebook, whether that was from an instructor, classmate (either directly or on Piazza), and/or online resources other than official Python documentation websites like docs.python.org or numpy.org. Alternatively, if you prefer, you may acknowledge assistance at the relevant point(s) in your code using a Python comment (#). Don't forget that you can receive extra credit from answering at least one question on Ed Discussion!

```python

```

*Acknowledge assistance here or in comments above your code:*





<!-- BEGIN QUESTION -->

# Question 1 - Numpy documentation (30 points)

There are many more functions and packages in Python than we have discussed in class. It is important for data analysts to be able to independently **find**, **learn about**, and **apply** useful functions.

The Numpy package contains many more functions than those we have covered in class. For this question, you will peruse the online Numpy package documentation and select **three** functions to learn more about and demonstrate (https://numpy.org/doc/stable/reference/routines.html). <br>
Also, don't forget to import numpy! You only have to do this once for your notebook.

For each Numpy function (10 points per function):


1. Describe what the function does. What value(s) does it return? <br>
2. Describe at least *two* **positional** arguments used in the function. What data types are required for them? What order do they go in? <br>
3. Describe at least *one* **optional** named argument used in the function. What data types are required for it? What is the default value?<br>
4. Provide the link to the documentation page(s) you used to answer these questions.<br>
5. Write a code snippet where you use the function and print the result.<br>

### Example: ```numpy.lcm```
1. The function numpy.lcm finds and returns the lowest common multiple of the absolute value of the inputs.<br>
2. The first two positional arguments used in numpy.lcm have to be either numpy arrays of integers or single scalar integers. An array or scalar of the lowest common multiple(s) of both arguments is returned by the function.<br>
3. One optional argument is "where" which indicates whether a universal function should be calculated at this point. The data type required is a Boolean. "where" = True is the default value.<br>
4. I used two references: https://numpy.org/doc/stable/reference/generated/numpy.lcm.html#numpy.lcm
and
https://www.educative.io/answers/what-is-numpylcm-in-python <br>


```python
#Example part 5: numpy.lcm. Do not change this code

import numpy as np #import numpy

lcm = np.lcm(45,6) #take the lowest common multiple of 45 and 6
print("The lowest common multiple of 45 and 6 is", lcm) #print

#Write your answers and code in the markdown cells below
```

_Type your answer here, replacing this text._

```python
## Part 5 Code for Function 1
```

```python
## Part 5 Code for Function 2
```

```python
## Part 5 Code for Function 3
```

<!-- END QUESTION -->

# Question #2 - Multi-diminsional array slicing (30 points)

### Instructions:
There is a secret message hidden in the following 3-D numpy array! We will decode this 3-D array by following the clues below and using our array indexing and slicing skills. <br>
_**Make sure to use the provided variables to store your final answer!**_ You may name intermediary variables whatever you like, as long as the final answer is stored and printed in the ones we provide in your starter code.

1. Find letters of words through indexing and slicing this array using the clues below to create an array of characters (letters). 
2. Turn resulting character arrays into lists using the numpy tolist() function.
3. Use the join() string function on each word list to form a single string for each word. 
4. Print the combined words at the end.

### Clues: 
Word 1: 5th layer, 3rd and 4th rows, and 2nd column.

Word 2: 1st to last layers, 4th row, 3rd column.

Word 3: 1st layer, 3rd row, 2nd to last columns.

Word 4: 4th layer, 3rd and 4th rows, 4th column.

Word 5: 3rd layer, 1st row, 1st to last columns
- For this clue, you will need to call the numpy ravel() function before you convert the character array into a list
- Read about what ravel() does [here.](https://numpy.org/doc/stable/reference/generated/numpy.ravel.html) 


### HINT
Example word: 1st layer, last row, 2nd to 4th columns
- This clue refers to the letters: ```['t','a','g']```

**Remember that Python indexing is non-inclusive and starts at 0, e.g. the 1st element is actually at index 0.**

```python
#Do not change this 3-D numpy string array
import numpy as np

str_array = np.array([[['a', 'c', 'V', 't', 'a'], #This is Layer 1
  ['M', 'h', 'c', 'r', 'b'], # Ex: This is Row 2
  ['T', 'n', 'a', 'm', 'e'], 
  ['g', 'f', 'c', 'A', 'a'],
  ['l', 't', 'a', 'g', 'h']],

 [['G', 'a', 'A', 'P', 'f'], #This is Layer 2
  ['x', 'q', 'q', 'h', 'b'],
  ['y', 'y', 'p', 'l', 'a'],
  ['a', 'a', 'a', 'a', 'U'],
  ['t', 't', 'a', 'K', 'y']],

 [['P', 'o', 's', 'c', 'o'], #This is Layer 3
  ['i', 'a', 'a', 'k', 'c'],
  ['r', 'a', 'a', 'f', 'k'],
  ['e', 'a', 't', 'a', 'r'],
  ['w', 'R', 'R', 'e', 's']],

 [['a', 'a', 'e', 'a', 'a'], #This is Layer 4
  ['t', 'b', 'r', 'a', 's'],
  ['e', 'w', 'a', 'i', 'a'],
  ['E', 'a', "'", 's', 's'],
  ['D', 'a', 'r', 'r', 'a']],

 [['E', 'a', 'a', 'r', 'r'], #This is Layer 5
  ['q', 'a', 'a', 'a', 'L'],
  ['j', 'M', 'a', 'y', 'Y'],
  ['j', 'y', 's', 'a', 'p'],
  ['e', 'a', 'a', 'r', 'a']]])


#Provide your code in the cell below
```

```python
#index 3-D array and turn resulting arrays into lists
word1_list = ...
word2_list = ...
word3_list = ...
word4_list = ...
word5_list = ...


#Use string function join to combine letters in to words
word1 = ...
word2 = ...
word3 = ...
word4 = ...
word5 = ...

#print full message
print(word1, word2, word3, word4, word5)
```

```python
grader.check("Question 2")
```

## Question 3 (30 points) - Datetimes, numpy, indexing, plotting  <font color='red'><i>Grade 30/30</i></font> 

![picture](https://drive.google.com/uc?export=view&id=1Vef5rWSHWuhFOO-jogERwF-wLMK62UE6)
*Image: Bathymetry of the seafloor around Axial Seamount (source: Wikipedia, originally from NOAA/PMEL).*

### **Axial Seamount eruption**

*Useful resources:* pre-lectures on list indexing, on *for* loops and *if* statements, on numpy, on datetime, and on basic plots

During the eruption of a submarine volcano, the seafloor deflates rapidly as magma is expelled into the ocean. Bottom pressure gauges use pressure to measure depth, so these gauges can detect this deflation. After an eruption, the seafloor slowly inflates and rises as magma accumulates in preparation for the next eruption.

Axial Seamount is an active submarine volcano offshore of the Oregon coast. It last erupted in 2015. Below, we provide a time series of monthly depths informed by bottom pressure data released by the [Ocean Observatories Initiative \(OOI\)](https://dataexplorer.oceanobservatories.org/#default-data/4) from August 15, 2014 to December 15, 2019, with measurements taken on the 15th day of each month. Times are specified in date strings.

## Instructions
Write your code and print your answers for Parts 1-3 in the specified cells below. <br>Make sure to write your *print()* statements for your answers at the bottom of each cell and _**use the provided variables to store your final answer!**_ You may name intermediary variables whatever you like, as long as the final answer is stored and printed in the ones we provide in your starter code. <br>You can choose whether to use lists or numpy arrays. Write code to answer the following questions:

1. The seafloor dropped more than 2.0 m during the 2015 eruption (AKA its depth increased/got deeper over 2m throughout the explosion). **This can be used to identify when the eruption happened**. Answer these questions using a *print()* statement, e.g. "Part 1a: The seafloor drop during the eruption was \<X> m." (10 points)
> **a.** Convert each date string in the "dates" list into a datetime and store them in a new list called "datetime_list". Print this list. 
>
> **b.** Think about how you can calculate the seafloor drop by using the known depths in the "depths" list. Calculate the seafloor drop during the eruption. Round to 1 decimal place (use function *round()*). Print the value of the seaflow drop.
>
> **c.** Identify the two measurement datetimes immediately before and after the eruption and save them into new variables. Print these dates.
>
> **d.** Find the seafloor depth measurement immediately before the eruption and save it as a new variable. Round to 1 decimal place. Print this variable.
>

2. The seafloor has been steadily inflating since the 2015 eruption (AKA its depth has been decreasing/becoming shallower since the eruption). Express answers to the following questions in meters and fractional years, respectively, each rounded to 1 decimal place: (5 points)
> **a.** How much has the seafloor inflated between the measurement immediately after the eruption and the most recent measurement in this data? Store the inflation difference in a variable and print it. 
>
> **b.** How much time has elapsed between these two measurements? Calculated the time difference and print it. 
>

3. Calculate the approximate rate of seafloor inflation using the values from Part 2. Express your answer in units of centimeters per week, rounded to 1 decimal place. Print your answer. (5 points)
- Rememeber that rate is in units of distance over time! 

4. Use Matplotlib to create a plot of the data with datetime on the x-axis and depth on the y-axis.  Shallower depths should point in the positive y-direction (Look how to do this in the [Matplotlib documentation](https://matplotlib.org/3.3.1/api/axes_api.html)). **Format and label your plot such that it looks as similar to the following plot as possible (this includes the title, axis labels, and grid)**: (10 points)

![picture](https://drive.google.com/uc?export=view&id=17LjK79LlBLUmqjkg6CDXjWhWNzTC8c8R)


```python
# Keep these starting lines of code – DO NOT CHANGE THIS!:
#
# Axial Seamount Central Caldera bottom pressure gauge data 
# (times in date strings; depths in meters)
dates = ["2014-08-15","2014-09-15","2014-10-15","2014-11-15","2014-12-15",
         "2015-01-15","2015-02-15","2015-03-15","2015-04-15","2015-05-15","2015-06-15","2015-07-15","2015-08-15","2015-09-15","2015-10-15","2015-11-15","2015-12-15",
         "2016-01-15","2016-02-15","2016-03-15","2016-04-15","2016-05-15","2016-06-15","2016-07-15","2016-08-15","2016-09-15","2016-10-15","2016-11-15","2016-12-15",
         "2017-01-15","2017-02-15","2017-03-15","2017-04-15","2017-05-15","2017-06-15","2017-07-15","2017-08-15","2017-09-15","2017-10-15","2017-11-15","2017-12-15",
         "2018-01-15","2018-02-15","2018-03-15","2018-04-15","2018-05-15","2018-06-15","2018-07-15","2018-08-15","2018-09-15","2018-10-15","2018-11-15","2018-12-15",
         "2019-01-15","2019-02-15","2019-03-15","2019-04-15","2019-05-15","2019-06-15","2019-07-15","2019-08-15","2019-09-15","2019-10-15","2019-11-15","2019-12-15"]
depths = [1510.24,1510.19,1510.14,1510.08,1510.04,1509.96,1509.93,1509.87,1509.83,1512.23,1512.11,1512.04,1511.96,1511.89,1511.79,1511.7,1511.64,1511.57,1511.54,1511.5,1511.47,1511.47,1511.4,1511.35,1511.33,1511.32,1511.31,1511.26,1511.21,1511.19,1511.14,1511.12,1511.05,1511.0,1511.06,1511.04,1511.01,1510.98,1510.93,1510.87,1510.87,1510.91,1510.86,1510.77,1510.75,1510.7,1510.67,1510.68,1510.64,1510.59,1510.56,1510.54,1510.45,1510.51,1510.41,1510.42,1510.39,1510.37,1510.4,1510.42,1510.41,1510.38,1510.33,1510.3,1510.26]

# Provide your answers in cells below:
```

## Part 1
The seafloor dropped more than 2.0 m during the 2015 eruption (AKA its depth increased/got deeper over 2m throughout the explosion). **This can be used to identify when the eruption happened**. Answer these questions using a *print()* statement, e.g. "Part 1a: The seafloor drop during the eruption was \<X> m." (10 points)

### Instructions
> **a.** Convert each date string in the "dates" list into a datetime and store them in a new list called "datetime_list". Print this list. 
>
> **b.** Think about how you can calculate the seafloor drop by using the known depths in the "depths" list. Calculate the seafloor drop during the eruption. Round to 1 decimal place (use function *round()*). Print the value of the seaflow drop.
>
> **c.** Identify the two measurement datetimes immediately before and after the eruption and save them into new variables. Think about how you would figure out the two months the eruption happend in between by using the seafloor drop value you calculated in Part B. Print these dates.
>
> **d.** Find the seafloor depth measurement immediately before the eruption and save it as a new variable. Round to 1 decimal place. Print this variable.
>

```python
## first, what do you need to import?
...

# Store your list of datetimes in here (part a)
datetime_list=...

...
# Store your seafloor drop (depth change) during the eruption in here (part b)
round_depth = ...

...
# Store your start and end dates of the eruption in here (part c)
start_date = ...
end_date = ...

...
# Store your depth immediately before the eruption in here (part d)
round_start_depth = ...
```

## Part 2
The seafloor has been steadily inflating since the 2015 eruption (AKA its depth has been decreasing/becoming shallower since the eruption). Express answers to the following questions in meters and fractional years, respectively, each rounded to 1 decimal place: (5 points)
> **a.** How much has the seafloor inflated between the measurement immediately after the eruption and the most recent measurement in this data? Store the inflation difference in a variable and print it. 
>
> **b.** How much time has elapsed between these two measurements? Calculated the time difference and print it. 
>

```python
...
# store the inflated difference in this variable (part a)
inflation_diff = ...

...
# store the elapsed time between measurements in this variable (part b)
elapsed_time = ...
```

### Part 3 
3. Calculate the approximate rate of seafloor inflation using the values from Part 2. Express your answer in units of centimeters per week, rounded to 1 decimal place. Print your answer. (5 points)
- Rememeber that rate is in units of distance over time! 

```python
...
# store rate of seafloor inflation in this variable
round_inflation_rate = ...
```

<!-- BEGIN QUESTION -->

### Part 4
Use Matplotlib to create a plot of the data with datetime on the x-axis and depth on the y-axis.  Shallower depths should point in the positive y-direction (Look how to do this in the [Matplotlib documentation](https://matplotlib.org/3.3.1/api/axes_api.html)). **Format and label your plot such that it looks as similar to the plot in the instructions as possible (this includes the title, axis labels, and grid)** (10 points)

```python
# Your code to make the plot for Part 4:
...
```

<!-- END QUESTION -->



Tests for Question 2 (points: 15)
```
>>> # word 1
>>> assert word1_list == str_array[4,2:4,1].tolist()
```
```
>>> # word 2
>>> assert word2_list == str_array[0:,3,2].tolist()
```
```
>>> # word 3
>>> assert word3_list == str_array[0,2,1:].tolist()
```
```
>>> # word 4
>>> assert word4_list == str_array[3,2:4,3].tolist()
```
```
>>> # word 5
>>> assert word5_list == np.ravel(str_array[2:4, 0:3,0]).tolist()
```


Tests for Question 3, Part 1 (points: 5)


Tests for Question 3, Part 2 (points: None)


Tests for Question 3, Part 3 (points: 2)
